# Lambda Function

Construct and plot the lambda function, defined as:
$$
\Lambda(q,x) := \sum_{i=0}^{n} |l_i(x)| 
$$

```python
# This is used for plots and numpy
%pylab inline 
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

```python
def lagrange_basis(xi, i):
    def func(x):
        assert i<len(xi) and i>=0, 'Out of range: 0 < i < len(xi)'
        ret = 1;
        for xj in xi[range(i)+range(i+1,len(xi))]:
            p = (x-xj)/(xi[i]-xj)
            ret *= p
        return ret
    return func
```

Test your implementation, plotting $l_i(x)$ for five equally spaced points.

<!-- grade_id: lagrange_basis, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-4-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: lambda_func_0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-5-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: lambda_func_0_test, points: 2, grade: true, solution: false -->
```python
from numpy.testing import *

assert_equal(lambda_function(linspace(0,1,3))(.25),1.25)
assert_almost_equal(lambda_function(linspace(0,1,5))(.3),1.2464,decimal=8)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-6-50203d367469> in <module>()
      1 from numpy.testing import *
      2 
----> 3 assert_equal(lambda_function(linspace(0,1,3))(.25),1.25)
      4 assert_almost_equal(lambda_function(linspace(0,1,5))(.3),1.2464,decimal=8)

NameError: name 'lambda_function' is not defined
```

Plot $\Lambda(q)(x)$ for equispaced nodes on the interval $[0,1]$ for 3, 5, and 7 equispaced nodes respectively. Comment the result.

<!-- grade_id: lambda_func_1_0, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-7-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: lambda_func_1_1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Chebyshev Nodes

Chebyshev nodes are defined as:
$$
x_k = cos\left(\frac{2k-1}{2n} \pi \right), k = 1,\ldots,n.
$$
Write a function to evaluate then for a given $n$.


<!-- grade_id: chebyshev_nodes_0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-8-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: chebyshev_nodes_0_test, points: 2, grade: true, solution: false -->
```python
assert_almost_equal(chebyshev_nodes(5),[  9.51056516e-01, 5.87785252e-01,   6.12323400e-17,  -5.87785252e-01,
  -9.51056516e-01],decimal=8)
assert_almost_equal(chebyshev_nodes(7),[  9.74927912e-01,   7.81831482e-01,   4.33883739e-01,   6.12323400e-17,
  -4.33883739e-01,  -7.81831482e-01,  -9.74927912e-01],decimal=8)

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-545121785039> in <module>()
----> 1 assert_almost_equal(chebyshev_nodes(5),[  9.51056516e-01, 5.87785252e-01,   6.12323400e-17,  -5.87785252e-01,
      2   -9.51056516e-01],decimal=8)
      3 assert_almost_equal(chebyshev_nodes(7),[  9.74927912e-01,   7.81831482e-01,   4.33883739e-01,   6.12323400e-17,
      4   -4.33883739e-01,  -7.81831482e-01,  -9.74927912e-01],decimal=8)

NameError: name 'chebyshev_nodes' is not defined
```

Plot the location of chebyshev nodes on the unit circle, and on the orizontal axis.

 - plot the positive sector of a cirumference with center $0$ and radius $1$.
 - plot the location of chebicev points on the cirumference, suggested plot option is `'*'`.
 - plot projection lines from the circle center to chebyshev nodes projected on the circumference.
 - plot projection lines from chebyshev nodes projected on the circumference to chebyshev nodes on the horizontal axis.
 - plot chebyshev nodes, suggested plot option is `'o'`.
 - make sure the axis are such that the resulting plot is truly a circle...

<!-- grade_id: chebyshev_nodes_1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-10-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

Plot $\Lambda(q)(x)$ for chebyshev nodes on the interval $[x_0,x_n]$ for 3, 5, and 7 nodes respectively. Comment the result.

<!-- grade_id: labda_on_cnodes_1, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-11-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: labda_on_cnodes_2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Define a function `function_norm` that returns $||f(x)||\infty$ with $x\in[a,b]$ (default to $[0,1]$). To keep the implementation simple, evaluate the given function over $2^n+1$ points and return the maximum value. Any smarter and correct implementation is welcome. $n$ should be selectable at run time, and it should be by default equal to 20 (rhoughly a million points)

<!-- grade_id: infty_norm_0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-12-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: infty_norm_0_test, points: 2, grade: true, solution: false -->
```python
print lambda_function(linspace(0,1,7))(.2)
print lambda_function(linspace(0,1,7))(linspace(0,1,2**2+1))
assert_almost_equal(function_norm(lambda_function(linspace(0,1,7)),0,1,10),4.54922445348, decimal=5)
assert_almost_equal(function_norm(lambda_function(linspace(0,1,5)),0,1,10),2.20780651848, decimal=5)

xi = chebyshev_nodes(7)
assert_almost_equal(function_norm(lambda_function(xi),xi[0],xi[-1]),1.78252988655, decimal=5)

xi = chebyshev_nodes(5)
assert_almost_equal(function_norm(lambda_function(xi),xi[0],xi[-1]),1.57016748842, decimal=5)
```

Output:
```
File "<ipython-input-13-ee65f74261a8>", line 1
    print lambda_function(linspace(0,1,7))(.2)
                        ^
SyntaxError: invalid syntax
```

Evaluate the $||\Lambda(q)||_{\infty}$ for degree from 1 to 7, for equispaced points, and for chebyshev nodes. Store the corresponding values in two different arrays `norm_l_eq`, and `norm_l_ch`.

<!-- grade_id: infty_norm_1, points: 2, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-14-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: infty_norm_1_test, points: 2, grade: true, solution: false -->
```python
assert_almost_equal(norm_l_eq,[1.,1.,1.25,1.63113031,2.2078244,3.10630116,4.54934179], decimal=5)
assert_almost_equal(norm_l_ch,[1.,1.,1.25,1.42987272,1.57016749,1.68514046,1.78252989], decimal=5)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-15-64e062523e9c> in <module>()
----> 1 assert_almost_equal(norm_l_eq,[1.,1.,1.25,1.63113031,2.2078244,3.10630116,4.54934179], decimal=5)
      2 assert_almost_equal(norm_l_ch,[1.,1.,1.25,1.42987272,1.57016749,1.68514046,1.78252989], decimal=5)

NameError: name 'norm_l_eq' is not defined
```

Plot the $||\Lambda(q,x)||_{\infty}$ for degree from 1 to 7, for equispaced points, and for chebyshev nodes.

<!-- grade_id: infty_norm_2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-16-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

# Bernstein Basis
Write a function that returns the ith Bernstein Basis of order $n$, using the standard definition: 
$$
B_{i,n}(t) = \binom{n}{i} t^i(1-t)^{n-i}
$$

<!-- grade_id: bern_0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-17-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_plot, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-18-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_0_test, points: 2, grade: true, solution: false -->
```python
expected = [[0.48999999999999994, 0.42, 0.09], [0.3429999999999999, 0.4409999999999999, 0.189, 0.026999999999999996], [0.24009999999999995, 0.4115999999999999, 0.2646, 0.07559999999999999, 0.0081], [0.16806999999999994, 0.3601499999999999, 0.3086999999999999, 0.13229999999999997, 0.028349999999999993, 0.0024299999999999994]]

for n in range(2,3):
    for i in range(0,n+1):
        b = bernstein_polynomial(n,i)
        assert_almost_equal(b(.3),expected[n-2][i],decimal=8)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-19-fd70133e9aef> in <module>()
      3 for n in range(2,3):
      4     for i in range(0,n+1):
----> 5         b = bernstein_polynomial(n,i)
      6         assert_almost_equal(b(.3),expected[n-2][i],decimal=8)

NameError: name 'bernstein_polynomial' is not defined
```

Implement the recursive definition of the bernstein polynomial:
$$
B_{k,n}(t) = (1-t)\cdot B_{k,n-1}(t)+t\cdot B_{k-1,n-1}(t)
$$

 - if $k < 0\ \mathrm{or}\ i > n$ then $B_{k,n} = 0$. 
 - if $k = 0\ \mathrm{and}\ n =0$ then $B_{k,n} = 1$. 

<!-- grade_id: bern_1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-20-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_1_test, points: 2, grade: true, solution: false -->
```python
expected = [[0.48999999999999994, 0.42, 0.09], [0.3429999999999999, 0.4409999999999999, 0.189, 0.026999999999999996], [0.24009999999999995, 0.4115999999999999, 0.2646, 0.07559999999999999, 0.0081], [0.16806999999999994, 0.3601499999999999, 0.3086999999999999, 0.13229999999999997, 0.028349999999999993, 0.0024299999999999994]]

for n in range(2,3):
    for i in range(0,n+1):
        b = bernstein_recursive(n,i)
        assert_almost_equal(b(.3),expected[n-2][i],decimal=8)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-21-9c4107493a57> in <module>()
      3 for n in range(2,3):
      4     for i in range(0,n+1):
----> 5         b = bernstein_recursive(n,i)
      6         assert_almost_equal(b(.3),expected[n-2][i],decimal=8)

NameError: name 'bernstein_recursive' is not defined
```

Implement power basis formula:
$$
B_{k,n}(t) = \sum_{i=k}^n(-1)^{i-k}\binom{n}{i}\binom{i}{k}t^i
$$

<!-- grade_id: bern_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-22-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_2_test, points: 2, grade: true, solution: false -->
```python
expected = [[0.48999999999999994, 0.42, 0.09], [0.3429999999999999, 0.4409999999999999, 0.189, 0.026999999999999996], [0.24009999999999995, 0.4115999999999999, 0.2646, 0.07559999999999999, 0.0081], [0.16806999999999994, 0.3601499999999999, 0.3086999999999999, 0.13229999999999997, 0.028349999999999993, 0.0024299999999999994]]

for n in range(2,3):
    for i in range(0,n+1):
        b = bernstein_power_basis(n,i)
        assert_almost_equal(b(.3),expected[n-2][i],decimal=8)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-23-10aa197231e3> in <module>()
      3 for n in range(2,3):
      4     for i in range(0,n+1):
----> 5         b = bernstein_power_basis(n,i)
      6         assert_almost_equal(b(.3),expected[n-2][i],decimal=8)

NameError: name 'bernstein_power_basis' is not defined
```

Plot all basis obtained through the "definition" implementation and through the "power basis" implementation for n=20,40,80,100. Give an explanation for what you observe.

<!-- grade_id: bern_3, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-24-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_3_bis, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Given the basis functions above, construct the Bernstein approximation of a continuous function, defined as:
$$
B_n(f)(t) := \sum_{i=0}^{n} B_{i,n} (t) * f(i/n) 
$$

<!-- grade_id: bern_func, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-25-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: bern_func_test, points: 2, grade: true, solution: false -->
```python
func = lambda t : sin(2*pi*t)
bi = bernstein_interpolation(3,func)
expected = [  0.00000000e+00,1.92450090e-01,-1.92450090e-01,-2.44929360e-16]
assert_almost_equal(bi(linspace(0,1,4)),expected,decimal=8)
bi = bernstein_interpolation(15,func)
expected = [  0.00000000e+00,6.50601790e-01,-6.50601790e-01,-2.44929360e-16]
assert_almost_equal(bi(linspace(0,1,4)),expected,decimal=8)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-26-8bec7f9c065c> in <module>()
      1 func = lambda t : sin(2*pi*t)
----> 2 bi = bernstein_interpolation(3,func)
      3 expected = [  0.00000000e+00,1.92450090e-01,-1.92450090e-01,-2.44929360e-16]
      4 assert_almost_equal(bi(linspace(0,1,4)),expected,decimal=8)
      5 bi = bernstein_interpolation(15,func)

NameError: name 'bernstein_interpolation' is not defined
```

Test your implementation, plotting the bernstein interpolation of $\sin(2\pi x)$ for different degrees:

<!-- grade_id: bernstein_interpolation_0, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-27-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

Given tolerance `tol` for the the function `func`, find the order $n$ that satisfies $\mathtt{tol}<||f(x)||_\infty$. Use the `function_norm` previously implemented.

<!-- grade_id: order_tol, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-28-15b94d1fa268> in <module>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: order_tol_test, points: 2, grade: true, solution: false -->
```python
func = lambda t : sin(2*pi*t)
assert_equal(find_order_given_tolerance(func,.05),77)

h = lambda x: abs(x-.5)
assert_equal(find_order_given_tolerance(h,.05),64)

g = lambda x: 1/(1+(10*x-5)**2)
assert_equal(find_order_given_tolerance(g,.1),187)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-29-97669fb97e0a> in <module>()
      1 func = lambda t : sin(2*pi*t)
----> 2 assert_equal(find_order_given_tolerance(func,.05),77)
      3 
      4 h = lambda x: abs(x-.5)
      5 assert_equal(find_order_given_tolerance(h,.05),64)

NameError: name 'find_order_given_tolerance' is not defined
```

```python

```

```python

```

```python

```